[1714790946] SERVICE ALERT: dev-dmz;Memory Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714791034] SERVICE ALERT: dev-dmz;Up Time;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714791066] SERVICE ALERT: dev-dmz;Memory Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714791079] HOST ALERT: prod-frontend;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1714791153] HOST ALERT: prod-frontend;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1714791154] SERVICE ALERT: dev-dmz;Up Time;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714791252] HOST ALERT: prod-frontend;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1714791263] SERVICE ALERT: dev-dmz;Memory Usage;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714791274] SERVICE ALERT: dev-dmz;Up Time;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714791326] HOST ALERT: prod-frontend;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1714791400] HOST ALERT: prod-frontend;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1714791474] HOST ALERT: prod-frontend;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1714791548] HOST ALERT: prod-frontend;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1714791622] HOST ALERT: prod-frontend;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1714791696] HOST ALERT: prod-frontend;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1714791770] HOST ALERT: prod-frontend;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1714792418] SERVICE ALERT: dev-backend;SWAP Usage;UNKNOWN;SOFT;1;NRPE: Command 'check_swap' not defined
[1714792439] SERVICE ALERT: dev-backend;SWAP Usage;OK;SOFT;2;Swap usage: 17%
[1714792442] SERVICE ALERT: dev-frontend;CPU Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714792562] SERVICE ALERT: dev-frontend;CPU Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714792638] SERVICE ALERT: dev-dmz;CPU Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714792670] SERVICE ALERT: dev-frontend;SWAP Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714792682] SERVICE ALERT: dev-frontend;CPU Usage;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714792758] SERVICE ALERT: dev-dmz;CPU Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714792790] SERVICE ALERT: dev-frontend;SWAP Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714792878] SERVICE ALERT: dev-dmz;CPU Usage;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714792888] SERVICE ALERT: dev-dmz;SWAP Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714792908] SERVICE ALERT: dev-frontend;Disk Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714792910] SERVICE ALERT: dev-frontend;SWAP Usage;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714793008] SERVICE ALERT: dev-dmz;SWAP Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714793028] SERVICE ALERT: dev-frontend;Disk Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714793127] SERVICE ALERT: dev-dmz;Disk Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714793128] SERVICE ALERT: dev-dmz;SWAP Usage;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714793148] SERVICE ALERT: dev-frontend;Disk Usage;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714793159] SERVICE ALERT: dev-backend;SWAP Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.67.63.94 port 5666: Connection refused
[1714793247] SERVICE ALERT: dev-dmz;Disk Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714793262] SERVICE ALERT: dev-backend;Disk Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.67.63.94 port 5666: Connection refused
[1714793279] SERVICE ALERT: dev-backend;SWAP Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.67.63.94 port 5666: Connection refused
[1714793367] SERVICE ALERT: dev-dmz;Disk Usage;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714793399] SERVICE ALERT: dev-backend;SWAP Usage;OK;SOFT;3;Swap usage: 17%
[1714793403] SERVICE ALERT: dev-backend;Disk Usage;OK;SOFT;2;OK - Disk usage: 64%
[1714793423] SERVICE ALERT: dev-dmz;Disk I/O Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714793543] SERVICE ALERT: dev-dmz;Disk I/O Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714793568] SERVICE ALERT: dev-frontend;Network Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714793590] SERVICE ALERT: dev-frontend;Disk I/O Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714793663] SERVICE ALERT: dev-dmz;Disk I/O Usage;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714793688] SERVICE ALERT: dev-frontend;Network Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714793710] SERVICE ALERT: dev-frontend;Disk I/O Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714793808] SERVICE ALERT: dev-frontend;Network Usage;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714793830] SERVICE ALERT: dev-frontend;Disk I/O Usage;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.124.73 port 5666: Connection refused
[1714793894] SERVICE ALERT: dev-dmz;Network Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714794014] SERVICE ALERT: dev-dmz;Network Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714794134] SERVICE ALERT: dev-dmz;Network Usage;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.65.95 port 5666: Connection refused
[1714794624] SERVICE ALERT: test-backend;RabbitMQ Overview;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.78.70 port 5666: Connection refused
[1714794744] SERVICE ALERT: test-backend;RabbitMQ Overview;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.78.70 port 5666: Connection refused
[1714794779] SERVICE ALERT: prod-backend;RabbitMQ Overview;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.83.110.50 port 5666: Connection refused
[1714794864] SERVICE ALERT: test-backend;RabbitMQ Overview;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.78.70 port 5666: Connection refused
[1714794899] SERVICE ALERT: prod-backend;RabbitMQ Overview;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.83.110.50 port 5666: Connection refused
[1714794949] SERVICE ALERT: prod-backend-backup;RabbitMQ Overview;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.96.174.95 port 5666: Connection refused
[1714794981] SERVICE ALERT: test-backend;RabbitMQ Queues;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.119.78.70 port 5666: Connection refused
[1714795019] SERVICE ALERT: prod-backend;RabbitMQ Overview;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.83.110.50 port 5666: Connection refused
[1714795069] SERVICE ALERT: prod-backend-backup;RabbitMQ Overview;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.96.174.95 port 5666: Connection refused
[1714795101] SERVICE ALERT: test-backend;RabbitMQ Queues;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.119.78.70 port 5666: Connection refused
[1714795174] SERVICE ALERT: prod-frontend;HTTP;CRITICAL;HARD;1;CRITICAL - Socket timeout
[1714795189] SERVICE ALERT: prod-backend-backup;RabbitMQ Overview;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.96.174.95 port 5666: Connection refused
[1714795221] SERVICE ALERT: test-backend;RabbitMQ Queues;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.119.78.70 port 5666: Connection refused
[1714795227] SERVICE ALERT: prod-backend;RabbitMQ Queues;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.83.110.50 port 5666: Connection refused
[1714795347] SERVICE ALERT: prod-backend;RabbitMQ Queues;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.83.110.50 port 5666: Connection refused
[1714795356] SERVICE ALERT: prod-backend-backup;RabbitMQ Queues;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 100.96.174.95 port 5666: Connection refused
[1714795467] SERVICE ALERT: prod-backend;RabbitMQ Queues;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.83.110.50 port 5666: Connection refused
[1714795476] SERVICE ALERT: prod-backend-backup;RabbitMQ Queues;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 100.96.174.95 port 5666: Connection refused
[1714795548] SERVICE ALERT: dev-frontend;Disk Usage;OK;HARD;1;OK - Disk usage: 75%
[1714795556] SERVICE ALERT: dev-frontend;CPU Usage;OK;HARD;1;CPU usage: 3.2%
[1714795568] SERVICE ALERT: dev-frontend;Disk I/O Usage;OK;HARD;1;Disk I/O - Read: 0 KB/s, Write: 3691 KB/s
[1714795584] SERVICE ALERT: dev-frontend;Memory Usage;OK;HARD;1;RAM usage: 37%
[1714795591] SERVICE ALERT: dev-frontend;Network Usage;OK;HARD;1;Network Traffic - Received: 16414 KB/s, Transmitted: 36599 KB/s
[1714795596] SERVICE ALERT: prod-backend-backup;RabbitMQ Queues;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 100.96.174.95 port 5666: Connection refused
[1714795597] SERVICE ALERT: dev-frontend;SWAP Usage;OK;HARD;1;Swap usage: 0%
[1714795607] SERVICE ALERT: dev-frontend;Up Time;OK;HARD;1;04:06:47 up 2 days,  3:02,  2 users,  load average: 0.01, 0.03, 0.00
[1714795688] SERVICE ALERT: dev-dmz;CPU Usage;OK;HARD;1;CPU usage: 0%
[1714795691] SERVICE ALERT: dev-dmz;Disk I/O Usage;OK;HARD;1;Disk I/O - Read: 0 KB/s, Write: 2263 KB/s
[1714795705] SERVICE ALERT: dev-dmz;Disk Usage;OK;HARD;1;OK - Disk usage: 64%
[1714795708] SERVICE ALERT: dev-dmz;Memory Usage;OK;HARD;1;RAM usage: 9%
[1714795711] SERVICE ALERT: dev-dmz;Network Usage;OK;HARD;1;Network Traffic - Received: 90 KB/s, Transmitted: 244 KB/s
[1714795713] SERVICE ALERT: dev-dmz;SWAP Usage;OK;HARD;1;Swap usage: 0%
[1714795716] SERVICE ALERT: dev-dmz;Up Time;OK;HARD;1;04:08:36 up 1 day,  2:37,  1 user,  load average: 0.36, 0.14, 0.05
[1714795916] SERVICE ALERT: test-backend;RabbitMQ Overview;OK;HARD;1;(No output returned from plugin)
[1714795918] SERVICE ALERT: test-backend;RabbitMQ Queues;OK;HARD;1;Total Messages: 0
[1714796217] SERVICE ALERT: prod-backend;RabbitMQ Overview;OK;HARD;1;(No output returned from plugin)
[1714796220] SERVICE ALERT: prod-backend;RabbitMQ Queues;OK;HARD;1;Total Messages: 0
[1714796364] SERVICE ALERT: prod-frontend;HTTP;OK;HARD;1;HTTP OK: HTTP/1.1 200 OK - 737 bytes in 0.003 second response time
[1714796368] HOST ALERT: prod-frontend;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.44 ms
[1714796631] SERVICE ALERT: prod-backend-backup;RabbitMQ Overview;OK;HARD;1;(No output returned from plugin)
[1714796634] SERVICE ALERT: prod-backend-backup;RabbitMQ Queues;OK;HARD;1;Total Messages: 0
[1714796680] SERVICE ALERT: dev-frontend;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1714796694] HOST ALERT: dev-frontend;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1714796749] SERVICE ALERT: dev-frontend;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1714796750] SERVICE ALERT: dev-frontend;HTTP;CRITICAL;HARD;4;CRITICAL - Socket timeout
[1714796758] SERVICE ALERT: dev-frontend;Disk Usage;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 10 seconds.
[1714796766] SERVICE ALERT: dev-frontend;CPU Usage;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 10 seconds.
[1714796768] HOST ALERT: dev-frontend;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1714796788] SERVICE ALERT: dev-frontend;Disk I/O Usage;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 10 seconds.
[1714796794] SERVICE ALERT: dev-frontend;Memory Usage;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 10 seconds.
[1714796801] SERVICE ALERT: dev-frontend;Network Usage;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 10 seconds.
[1714796807] SERVICE ALERT: dev-frontend;SWAP Usage;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 10 seconds.
[1714796817] SERVICE ALERT: dev-frontend;Up Time;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 10 seconds.
[1714796842] HOST ALERT: dev-frontend;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1714796916] HOST ALERT: dev-frontend;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1714796980] HOST ALERT: dev-frontend;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.89 ms
[1714797042] SERVICE ALERT: dev-frontend;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.12 ms
[1714797149] SERVICE FLAPPING ALERT: dev-frontend;CPU Usage;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1714797149] SERVICE ALERT: dev-frontend;CPU Usage;OK;HARD;1;CPU usage: 0%
[1714797153] SERVICE FLAPPING ALERT: dev-frontend;Disk I/O Usage;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1714797153] SERVICE ALERT: dev-frontend;Disk I/O Usage;OK;HARD;1;Disk I/O - Read: 0 KB/s, Write: 1455 KB/s
[1714797156] SERVICE FLAPPING ALERT: dev-frontend;Disk Usage;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1714797156] SERVICE ALERT: dev-frontend;Disk Usage;OK;HARD;1;OK - Disk usage: 76%
[1714797158] SERVICE ALERT: dev-frontend;HTTP;OK;HARD;1;HTTP OK: HTTP/1.1 200 OK - 722 bytes in 0.005 second response time
[1714797161] SERVICE ALERT: dev-frontend;Memory Usage;OK;HARD;1;RAM usage: 8%
[1714797164] SERVICE FLAPPING ALERT: dev-frontend;Network Usage;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1714797164] SERVICE ALERT: dev-frontend;Network Usage;OK;HARD;1;Network Traffic - Received: 16 KB/s, Transmitted: 10 KB/s
[1714797167] SERVICE FLAPPING ALERT: dev-frontend;SWAP Usage;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1714797167] SERVICE ALERT: dev-frontend;SWAP Usage;OK;HARD;1;Swap usage: 0%
[1714797169] SERVICE ALERT: dev-frontend;Up Time;OK;HARD;1;04:32:49 up 3 min,  0 users,  load average: 0.00, 0.00, 0.00
[1714798310] SERVICE FLAPPING ALERT: dev-frontend;CPU Usage;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1714798310] SERVICE FLAPPING ALERT: dev-frontend;Disk I/O Usage;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1714798310] SERVICE FLAPPING ALERT: dev-frontend;Disk Usage;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1714798310] SERVICE FLAPPING ALERT: dev-frontend;Network Usage;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1714798310] SERVICE FLAPPING ALERT: dev-frontend;SWAP Usage;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1714798860] SERVICE FLAPPING ALERT: dev-frontend;CPU Usage;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1714798860] SERVICE FLAPPING ALERT: dev-frontend;Disk I/O Usage;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1714798860] SERVICE FLAPPING ALERT: dev-frontend;Disk Usage;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1714798860] SERVICE FLAPPING ALERT: dev-frontend;Network Usage;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1714798860] SERVICE FLAPPING ALERT: dev-frontend;SWAP Usage;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)